
# Project 2 Validation Notebook

You cannot write any code in this notebook – to actually work on the project, open `project.py` and `project.ipynb`.

The only purpose of this notebook is to give you a blank copy of `project.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `project.py`.

**Before submitting Project 2, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `project.ipynb` that is not implemented correctly in `project.py`, or it could be that a function in `project.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("project.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

```python
import numpy as np
import pandas as pd
import os

from IPython.display import clear_output

import pathlib
import urllib.request
```

```python
from project import *
```

```python
from ipywidgets import widgets, interact, FloatSlider
import plotly.express as px
```

```python
vacs = pd.read_csv(os.path.join('data', 'covid_vaccinations.csv'))
vacs
```

Output:
```
Date      Country_Region   Doses_admin  People_at_least_one_dose
0       2020-12-29             Austria  2.123000e+03                      2123
1       2020-12-29             Bahrain  5.501400e+04                     55014
2       2020-12-29             Belarus  0.000000e+00                         0
3       2020-12-29             Belgium  3.400000e+02                       340
4       2020-12-29              Canada  5.907900e+04                     59078
...            ...                 ...           ...                       ...
133517  2023-01-20  West Bank and Gaza  3.748571e+06                   2012767
133518  2023-01-20               World  1.325698e+10                5514130774
133519  2023-01-20               Yemen  1.242982e+06                   1006274
133520  2023-01-20              Zambia  1.361498e+07                  10959569
133521  2023-01-20            Zimbabwe  1.269485e+07                   6558577

[133522 rows x 4 columns]
```

```python
DEFAULT = 'US'

def plot_cases(country):
    country_only = vacs[vacs['Country_Region'] == country]
    fig = px.line(country_only, x='Date', y='Doses_admin', title=f"'Doses_admin' column for {country}")
    fig.show()
    

dropdown_cases = widgets.Dropdown(options=np.sort(vacs['Country_Region'].unique()), value=DEFAULT)

def dropdown_cases_handler(change):
    if change['name'] == 'value' and (change['new'] != change['old']):
        clear_output()
        display(dropdown_cases)
        plot_cases(change['new'])
        
display(dropdown_cases)
plot_cases(DEFAULT)
dropdown_cases.observe(dropdown_cases_handler)
```

Output:
```
Dropdown(index=179, options=('Afghanistan', 'Albania', 'Algeria', 'Andorra', 'Angola', 'Antigua and Barbuda', …
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
mono = monotonic_violations_by_country(vacs)
mono
```

Output:
```
Doses_admin_monotonic  People_at_least_one_dose_monotonic
Country_Region                                                               
Afghanistan                             0                                   0
Albania                                 0                                   0
Algeria                                 0                                   0
Andorra                                 0                                   0
Angola                                  0                                   0
...                                   ...                                 ...
West Bank and Gaza                      0                                   0
World                                   2                                   6
Yemen                                   1                                   0
Zambia                                  0                                   0
Zimbabwe                                0                                   0

[194 
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
tots = robust_totals(vacs)
tots
```

Output:
```
Doses_admin  People_at_least_one_dose
Country_Region                                            
Afghanistan         1.244987e+07              1.145969e+07
Albania             3.025728e+06              1.342243e+06
Algeria             1.526744e+07              7.840131e+06
Andorra             1.543670e+05              5.790100e+04
Angola              2.370105e+07              1.499013e+07
...                          ...                       ...
West Bank and Gaza  3.748571e+06              2.012767e+06
World               1.319737e+10              5.503074e+09
Yemen               1.239152e+06              1.003208e+06
Zambia              1.358138e+07              1.093767e+07
Zimbabwe            1.269485e+07              6.558577e+06

[194 rows x 2 columns]
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

```python
pops_raw = pd.read_csv(os.path.join('data', 'populations.csv'))
pops_raw
```

Output:
```
Country (or dependency)  Population in 2023 World Percentage  \
0                     India         1428627.663           17.85%   
1                     China         1425671.352           17.81%   
2             United States          339996.563            4.25%   
3                 Indonesia          277534.122            3.47%   
4                  Pakistan          240485.658            3.00%   
..                      ...                 ...              ...   
200              San Marino              33.642            0.00%   
201                   Palau              18.058            0.00%   
202                   Nauru              12.780            0.00%   
203                  Tuvalu              11.396            0.00%   
204            Vatican City               0.518            0.00%   

        Area (Km²) Density (P/Km²)  ISO  
0    3_287_590 Km²    480.5033/Km²  IND  
1    9,706,961 Km²    151.2696/Km²  CHN  
2    9,372,610 Km²     37.1686/Km²  USA  
3    1,904,569 Km² 
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
pops = fix_dtypes(pops_raw)
pops
```

Output:
```
Country (or dependency)  Population in 2023  World Percentage  Area (Km²)  \
0                     India          1428627663            0.1785     3287590   
1                     China          1425671352            0.1781     9706961   
2             United States           339996563            0.0425     9372610   
3                 Indonesia           277534122            0.0347     1904569   
4                  Pakistan           240485658            0.0300      881912   
..                      ...                 ...               ...         ...   
200              San Marino               33642            0.0000          61   
201                   Palau               18058            0.0000         459   
202                   Nauru               12780            0.0000          21   
203                  Tuvalu               11396            0.0000          26   
204            Vatican City                 518            0.0000           1   

     Density (P/Km²)  ISO  
0  
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
missing = missing_in_pops(tots, pops)
missing
```

Output:
```
{'Burma',
 'Cabo Verde',
 'Congo (Brazzaville)',
 'Congo (Kinshasa)',
 "Cote d'Ivoire",
 'Czechia',
 'Korea, South',
 'Kosovo',
 'US',
 'West Bank and Gaza',
 'World'}
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
pops_fixed = fix_names(pops)
pops_fixed
```

Output:
```
Country (or dependency)  Population in 2023  World Percentage  Area (Km²)  \
0                     India          1428627663            0.1785     3287590   
1                     China          1425671352            0.1781     9706961   
2                        US           339996563            0.0425     9372610   
3                 Indonesia           277534122            0.0347     1904569   
4                  Pakistan           240485658            0.0300      881912   
..                      ...                 ...               ...         ...   
200              San Marino               33642            0.0000          61   
201                   Palau               18058            0.0000         459   
202                   Nauru               12780            0.0000          21   
203                  Tuvalu               11396            0.0000          26   
204            Vatican City                 518            0.0000           1   

     Density (P/Km²)  ISO  
0  
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
fig = draw_choropleth(tots, pops_fixed)
fig.show()
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

```python
# Download Israeli COVID vaccinations data from the ☁️
if not pathlib.Path(os.path.join('data', 'israel.csv')).exists():
    urllib.request.urlretrieve(
        'https://f000.backblazeb2.com/file/dsc-data/covid-israel/israel.csv',
        os.path.join('data', 'israel.csv')
    )
```

```python
israel_raw = pd.read_csv(os.path.join('data', 'israel.csv'))
israel_raw
```

Output:
```
Age  Vaccinated  Severe Sickness
0        23           1                0
1        60           1                0
2        43           1                0
3        40           1                0
4        71           1                0
...      ..         ...              ...
6937537  14           1                0
6937538  30           1                0
6937539  18           0                0
6937540  45           1                0
6937541  27           1                0

[6937542 rows x 3 columns]
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
israel = clean_israel_data(israel_raw)
israel
```

Output:
```
Age  Vaccinated  Severe Sickness
0        23.0        True            False
1        60.0        True            False
2        43.0        True            False
3        40.0        True            False
4        71.0        True            False
...       ...         ...              ...
6937537  14.0        True            False
6937538  30.0        True            False
6937539  18.0       False            False
6937540  45.0        True            False
6937541  27.0        True            False

[6937542 rows x 3 columns]
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

```python
israel_raw.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 6937542 entries, 0 to 6937541
Data columns (total 3 columns):
 #   Column           Dtype 
---  ------           ----- 
 0   Age              object
 1   Vaccinated       int64 
 2   Severe Sickness  int64 
dtypes: int64(2), object(1)
memory usage: 158.8+ MB
```

```python
israel.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 6937542 entries, 0 to 6937541
Data columns (total 3 columns):
 #   Column           Dtype  
---  ------           -----  
 0   Age              float64
 1   Vaccinated       bool   
 2   Severe Sickness  bool   
dtypes: bool(2), float64(1)
memory usage: 66.2 MB
```

```python
israel.assign(null_age=israel['Age'].isna()).groupby('null_age').mean()
```

Output:
```
Age  Vaccinated  Severe Sickness
null_age                                     
False     41.036    0.812204         0.000074
True         NaN    0.805164         0.000000
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
# expect this cell to take a while to run!
test_stats_vax, test_stats_sick = mcar_permutation_tests(israel, 100)
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

```python
vax = israel.dropna()
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
eff_overall = effectiveness(vax)
eff_overall
```

Output:
```
0.6747820515830074
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

```python
AGE_GROUPS = [
    '12-15',
    '16-19',
    '20-29',
    '30-39',
    '40-49',
    '50-59',
    '60-69',
    '70-79',
    '80-89',
    '90-'
]
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
s_eff = stratified_effectiveness(vax)
s_eff
```

Output:
```
Group
12-15    1.000000
16-19    1.000000
20-29    1.000000
30-39    0.966591
40-49    0.941211
50-59    0.926886
60-69    0.886400
70-79    0.895772
80-89    0.810277
90-      0.924404
dtype: float64
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
eff_example = effectiveness_calculator(
    young_vaccinated_prop=0.01,
    old_vaccinated_prop=0.99,
    young_risk_vaccinated=0.01,
    young_risk_unvaccinated=0.20,
    old_risk_vaccinated=0.10,
    old_risk_unvaccinated=0.50
)
eff_example
```

Output:
```
{'Young': 0.95, 'Old': 0.8, 'Overall': 0.5118226600985222}
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

```python
# The sliders below are in the same order as the arguments below,
# in case some of the slider labels are cut off.
interact(effectiveness_calculator,
                 young_vaccinated_prop=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.01),
                 old_vaccinated_prop=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.99),
                 young_risk_vaccinated=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.01),
                 young_risk_unvaccinated=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.20),
                 old_risk_vaccinated=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.10),
                 old_risk_unvaccinated=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.50));
```

Output:
```
interactive(children=(FloatSlider(value=0.01, description='young_vaccinated_prop', max=0.99, min=0.01, step=0.…
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
eff_extreme = effectiveness_calculator(**extreme_example())
eff_extreme
```

Output:
```
{'Young': 0.95, 'Old': 0.8305084745762712, 'Overall': -0.11395101171459}
```

```python
grader.check("q11")
```

Output:
```
q11 results: All test cases passed!
```

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q10 results: All test cases passed!

q11 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!

q5 results: All test cases passed!

q6 results: All test cases passed!

q7 results: All test cases passed!

q8 results: All test cases passed!

q9 results: All test cases passed!
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `project.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> count_monotonic(np.array([3, 6, 6, 2, 5, 8])) == 1
True
```
```
>>> isinstance(mono, pd.DataFrame)
True
```
```
>>> mono.shape == (194, 2) and set(mono.columns) == {'Doses_admin_monotonic', 'People_at_least_one_dose_monotonic'}
True
```
```
>>> mono.sum().sum() > 0
True
```


Tests for q10 (points: None)
```
>>> isinstance(eff_example, dict)
True
```
```
>>> eff_example.keys() == {'Overall', 'Young', 'Old'}
True
```
```
>>> test_eff = effectiveness_calculator(young_vaccinated_prop=0.5, old_vaccinated_prop=0.5, young_risk_vaccinated=0.01, young_risk_unvaccinated=0.20, old_risk_vaccinated=0.01, old_risk_unvaccinated=0.20)
>>> test_eff['Overall'] == test_eff['Young'] == test_eff['Old'] == 0.95
True
```
```
>>>
```


Tests for q11 (points: None)
```
>>> isinstance(eff_extreme, dict)
True
```
```
>>> keys = {'old_risk_unvaccinated', 'young_risk_vaccinated', 'young_vaccinated_prop', 'old_risk_vaccinated', 'old_vaccinated_prop', 'young_risk_unvaccinated'}
>>> extreme_example().keys() == keys
True
```
```
>>> eff_extreme.keys() == {'Overall', 'Young', 'Old'}
True
```


Tests for q2 (points: None)
```
>>> isinstance(tots, pd.DataFrame)
True
```
```
>>> tots.shape == (194, 2)
True
```
```
>>> len(set(tots.columns).difference({'Doses_admin', 'People_at_least_one_dose'}))==0
True
```
```
>>> list(tots.dtypes) == [float]*2
True
```


Tests for q3 (points: None)
```
>>> isinstance(pops, pd.DataFrame) and pops.shape == (205, 6)
True
```
```
>>> # pops.loc[pops['Country (or dependency)'] == 'Japan', 'Population in 2023'].iloc[0] == 123294513
>>> 
>>> len(set(pops.columns).difference({'Area (Km²)',
...  'Country (or dependency)',
...  'Density (P/Km²)',
...  'ISO',
...  'Population in 2023',
...  'World Percentage'})) == 0
True
```


Tests for q4 (points: None)
```
>>> isinstance(missing, set)
True
```
```
>>> ex_tots = pd.DataFrame({
...             'Doses_admin': [1, 2, 3],
...             'People_partially_vaccinated': [1, 2, 3],
...             'People_fully_vaccinated': [1, 2, 3]
...         },
...         index = ['China', 'Angola', 'Republic of Data Science']
...     )
>>> missing_in_pops(ex_tots, pops) == {'Republic of Data Science'}
True
```
```
>>> isinstance(pops_fixed, pd.DataFrame)
True
```
```
>>> pops_fixed.shape == (205, 6)
True
```
```
>>> 'Burma' in pops_fixed['Country (or dependency)'].values
True
```
```
>>> not 'Myanmar' in pops_fixed['Country (or dependency)'].values
True
```


Tests for q5 (points: None)
```
>>> str(type(fig)) == "<class 'plotly.graph_objs._figure.Figure'>" and fig.data[0].type == 'choropleth'
True
```
```
>>> 'Cameroon' in fig.data[0]['hovertext']
True
```
failure_message: Ensuring that the `hover_name` argument is set correctly (make sure it's not set to the 'ISO' column, and make sure you didn't set the `hover_data` argument instead)
```
>>> fig.layout['coloraxis']['colorbar']['title']['text'] == 'Doses Per Person'
True
```
failure_message: Ensures the legend is set correctly
```
>>> 'covid vaccine doses per person' in fig.layout['title']['text'].lower()
True
```
failure_message: Ensures the title is set correctly
```
>>> colors = [c[1] for c in fig.layout['coloraxis']['colorscale']]
>>> default_colorscale = ['#0d0887', '#46039f', '#7201a8', '#9c179e', '#bd3786', '#d8576b', '#ed7953', '#fb9f3a', '#fdca26', '#f0f921']
>>> colors != default_colorscale
True
```
failure_message: Ensuring changes were made to the colorscale
```
>>> plotted_data = pd.Series(fig.data[0]['z'], index=fig.data[0]['locations'])
>>> np.isclose(plotted_data['USA'], 1.9496818625192982, atol=0.5)
True
```
failure_message: Plotted the correct data for the USA. If this fails, and your plot looks right, make sure that you only set the `locations` argument, and not the `locationmode` argument.


Tests for q6 (points: None)
```
>>> isinstance(israel, pd.DataFrame)
True
```
```
>>> israel.shape == (6937542, 3)
True
```


Tests for q7 (points: None)
```
>>> isinstance(test_stats_vax, np.ndarray) and isinstance(test_stats_sick, np.ndarray)
True
```
```
>>> len(test_stats_vax) == len(test_stats_sick) == 100
True
```
```
>>> missingness_type() in {1, 2, 3, 4}
True
```


Tests for q8 (points: None)
```
>>> isinstance(eff_overall, float)
True
```
```
>>> 0 <= eff_overall <= 1
True
```
```
>>> example_vax = pd.DataFrame({'Age': [15, 20, 25, 30, 35, 40], 
...                             'Vaccinated': [True, True, True, False, False, False], 
...                             'Severe Sickness': [True, False, False, False, True, True]})
>>> effectiveness(example_vax) == 0.5
True
```


Tests for q9 (points: None)
```
>>> isinstance(s_eff, pd.Series) 
True
```
```
>>> len(s_eff) == 10
True
```
```
>>> israel_raw = pd.read_csv(os.path.join('data', 'israel-subset.csv'))
>>> vax_subset = clean_israel_data(israel_raw).dropna()
>>> stratified_effectiveness(vax_subset).index[0] == '12-15'
True
```
```
>>> israel_raw = pd.read_csv(os.path.join('data', 'israel-subset.csv'))
>>> vax_subset = clean_israel_data(israel_raw).dropna()
>>> len(stratified_effectiveness(vax_subset)) == 10
True
```
