
# Lab 2 Validation Notebook

You cannot write any code in this notebook – to actually work on the lab, open `lab.py` and `lab.ipynb`.

The only purpose of this notebook is to give you a blank copy of `lab.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `lab.py`.

**Before submitting Lab 2, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `lab.ipynb` that is not implemented correctly in `lab.py`, or it could be that a function in `lab.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
import os
```

```python
# don't change this cell -- it is needed for the tests to work
scores_fp = os.path.join('data', 'scores.csv')
scores = data_load(scores_fp)
passfail = pass_fail(scores.copy())
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

```python
# don't change this cell -- it is needed for the tests to work
medscore = med_score(passfail.copy())
highest = highest_score_name(passfail)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

```python
# don't change this cell -- it is needed for the tests to work
idxdup = idx_dup()
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

```python
# don't change this cell -- it is needed for the tests to work
trick_ans = trick_bool()
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

```python
# don't change this cell -- it is needed for the tests to work
pop_data = np.random.choice(range(10), size=(100, 4))
df_pop = pd.DataFrame(pop_data, columns='A B C D'.split())
out_pop = population_stats(df_pop)
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

```python
# don't change this cell -- it is needed for the tests to work
common_data = np.random.choice(range(10), size=(100, 2))
common_df = pd.DataFrame(common_data, columns='A B'.split())
common_out = most_common(common_df, N=3)
common_out
```

Output:
```
2    15
1    14
6    12
Name: A, dtype: int64
9    16
8    13
3    12
Name: B, dtype: int64
   A_values  A_counts  B_values  B_counts
0         2        15         9        16
1         1        14         8        13
2         6        12         3        12
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

```python
# don't change this cell -- it is needed for the tests to work
super_fp = os.path.join('data', 'superheroes_powers.csv')
powers = pd.read_csv(super_fp)
super_out = super_hero_powers(powers)
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

```python
# don't change this cell -- it is needed for the tests to work
superheroes_fp = os.path.join('data', 'superheroes.csv')
heroes = pd.read_csv(superheroes_fp, index_col=0)
clean_out = clean_heroes(heroes)
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

```python
clean_out.head(10)
```

Output:
```
name  Gender Eye color               Race Hair color  Height  \
0         A-Bomb    Male    yellow              Human    No Hair   203.0   
1     Abe Sapien    Male      blue      Icthyo Sapien    No Hair   191.0   
2       Abin Sur    Male      blue            Ungaran    No Hair   185.0   
3    Abomination    Male     green  Human / Radiation    No Hair   203.0   
4        Abraxas    Male      blue      Cosmic Entity      Black     NaN   
5  Absorbing Man    Male      blue              Human    No Hair   193.0   
6    Adam Monroe    Male      blue                NaN      Blond     NaN   
7   Adam Strange    Male      blue              Human      Blond   185.0   
8       Agent 13  Female      blue                NaN      Blond   173.0   
9      Agent Bob    Male     brown              Human      Brown   178.0   

           Publisher Skin color Alignment  Weight  
0      Marvel Comics        NaN      good   441.0  
1  Dark Horse Comics       blue      good    65.0  
2          DC Comic
```

```python
# don't change this cell -- it is needed for the tests to work
stats_out = super_hero_stats()
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q10 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!

q5 results: All test cases passed!

q6 results: All test cases passed!

q7 results: All test cases passed!

q8 results: All test cases passed!

q9 results: All test cases passed!
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `lab.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> isinstance(scores, pd.DataFrame)
True
```
```
>>> list(scores.columns) == ['attempts', 'highest_score']
True
```
```
>>> not isinstance(scores.index[0], int)
True
```
```
>>> isinstance(passfail, pd.DataFrame)
True
```
```
>>> len(passfail.columns) == 3
True
```
```
>>> passfail.loc["Julia", "pass"] == 'Yes'
True
```


Tests for q10 (points: None)
```
>>> isinstance(stats_out[0], str)
True
```
```
>>> isinstance(stats_out[1], str)
True
```
```
>>> stats_out[2] in ['good', 'bad']
True
```
```
>>> stats_out[3] in ['Marvel Comics', 'DC Comics']
True
```
```
>>> isinstance(stats_out[4], str)
True
```
```
>>> isinstance(stats_out[5], str)
True
```


Tests for q2 (points: None)
```
>>> isinstance(medscore, float)
True
```
```
>>> 89 < medscore < 91
True
```
```
>>> isinstance(highest, tuple)
True
```
```
>>> len(highest[1]) == 3
True
```


Tests for q3 (points: None)
```
>>> isinstance(idxdup, int)
True
```
```
>>> 1 <= idxdup <= 6
True
```


Tests for q4 (points: None)
```
>>> ans =  trick_me()
>>> ans == 1 or ans == 2 or ans == 3
True
```
```
>>> isinstance(trick_ans, list)
True
```
```
>>> isinstance(trick_ans[1], int)
True
```


Tests for q5 (points: None)
```
>>> change(1.0) == 1.0
True
```
```
>>> change(np.NaN) == 'MISSING'
True
```
```
>>> df_with_nans = pd.DataFrame([[0, 1, np.NaN], [np.NaN, np.NaN, np.NaN], [1, 2, 3]])
>>> A = correct_replacement(df_with_nans)
>>> ((A.values == 'MISSING').sum() == 4) & (A is not df_with_nans)
True
```
```
>>> ans =  missing_ser()
>>> ans == 1 or ans == 2 or ans == 3
True
```
```
>>> df_with_nans = pd.DataFrame([[0, 1, np.NaN], [np.NaN, np.NaN, np.NaN], [1, 2, 3]])
>>> fill_ser(df_with_nans)
>>> (df_with_nans.values == 'MISSING').sum() == 4
True
```


Tests for q6 (points: None)
```
>>> out_pop.index.tolist() == ['A', 'B', 'C', 'D']
True
```
```
>>> cols = ['num_nonnull', 'prop_nonnull', 'num_distinct', 'prop_distinct']
>>> out_pop.columns.tolist() == cols
True
```
```
>>> (out_pop['num_distinct'] <= 10).all()
True
```
```
>>> (out_pop['prop_nonnull'] == 1.0).all()
True
```


Tests for q7 (points: None)
```
>>> common_out.index.tolist() == [0, 1, 2]
True
```
```
>>> common_out.columns.tolist() == ['A_values', 'A_counts', 'B_values', 'B_counts']
True
```
```
>>> common_out['A_values'].isin(range(10)).all()
True
```


Tests for q8 (points: None)
```
>>> isinstance(super_out, list)
True
```
```
>>> len(super_out) == 3
True
```
```
>>> all([isinstance(x, str) for x in super_out])
True
```


Tests for q9 (points: None)
```
>>> clean_out['Skin color'].isnull().any()
True
```
```
>>> clean_out['Weight'].isnull().any()
True
```
